# Homework 4: Testing a Model

Read Chapter 11.3 of the textbook before starting.

This homework is one complete hypothesis test, start to finish.

**This homework has hidden tests on it.** That means even though the tests may
say 100% passed, it doesn't mean your final grade will be 100%. We will be
running more hidden tests for correctness once everyone turns in the homework.


Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](https://ulwazi.wits.ac.za/courses/89081/pages/detailed-python-reference-sheet-python-cheat-sheet-2?module_item_id=1200407): Cheat sheet of helpful array & table methods used!

**Recommended Readings**: 
* [Sampling and Empirical Distributions](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html)
* [A/B Testing](https://inferentialthinking.com/chapters/12/1/AB_Testing.html)

```python
# Run this cell first to set up the notebook -- the install takes about a minute in the browser
# Please don't change it.
%pip install -q urllib3<2.0 otter-grader datascience ipywidgets

try:
    import pyodide_http
    pyodide_http.patch_all()
except ImportError:
    pass

import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw4.ipynb")
```

---
---

## 1. Vaccinations Across The Nation

A vaccination clinic has two types of vaccines against a disease. Each person who comes in to be vaccinated gets either Vaccine 1 or Vaccine 2. One week, everyone who came in on Monday, Wednesday, and Friday was given Vaccine 1. Everyone who came in on Tuesday and Thursday was given Vaccine 2. The clinic is closed on weekends.

Doctor DeNero at the clinic said, "Oh wow, the distribution of vaccines is like tossing a coin that lands heads with probability $\frac{3}{5}$. If the coin lands on heads, you get Vaccine 1 and if the coin lands on tails, you get Vaccine 2."

But Doctor Sahai said, "No, it's not. We're not doing anything like tossing a (biased) coin."

That week, the clinic gave Vaccine 1 to 211 people and Vaccine 2 to 107 people. Conduct a test of hypotheses to see which doctor's position is better supported by the data.

---

**Question 1.1.** Given the information above, what was the sample size for the data, and what was the percentage of people who got **Vaccine 1?** **(4 points)**

*Note*: Your percent should be a number between 0 and 100, not a proportion between 0 and 1.


```python
sample_size = ...
percent_V1 = ...

print(f"Sample Size: {sample_size}")
print(f"Vaccine 1 Percent: {percent_V1}")
```

```python
grader.check("q1_1")
```

---

**Question 1.2.** One of the test statistics below is appropriate for testing these hypotheses. Assign the variable `valid_test_stat` to the number corresponding to the correct test statistic. **(4 points)**

_Hint:_ Recall that large values of the test statistic should favor the alternative hypothesis.

1. percent of heads - 50
2. |percent of heads - 50|
3. percent of heads - 60
4. |percent of heads - 60|


```python
valid_test_stat = ...
valid_test_stat
```

```python
grader.check("q1_2")
```

---

**Question 1.3.** Using your answer from Questions 1.1 and 1.2, find the observed value of the test statistic and assign it to the variable `observed_statistic`. Recall that the observed statistic is the test statistic value that was observed in the real life data. **(4 points)**


```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q1_3")
```

---

**Question 1.4.** Simulate 10,000 values of the test statistic under the null hypothesis — that is, under Dr. DeNero's claim that each person's vaccine is decided like a biased coin landing heads with probability $\frac{3}{5}$. **(8 points)**

Start by defining a function that simulates **one** value of the statistic. It should use `sample_proportions`, and it will need `sample_size` from Question 1.1. Then write a `for` loop that collects 10,000 values in an array called `simulated_statistics`.

*Note:* `sample_proportions(n, distribution)` returns the **proportion** of the sample in each category, so multiply by 100 to get a percentage.

*Note:* Your statistic must be computed the same way as the observed one in Question 1.3, or the comparison in Question 1.5 will be meaningless.

```python
def one_simulated_statistic():
    ...

# Run the this cell a few times to see how the simulated statistic changes
one_simulated_statistic()
```

```python
num_simulations = 10000

simulated_statistics = ...
for ... in ...:
    ...
```

```python
grader.check("q1_4")
```

```python
# Run this cell to produce a histogram of the simulated statistics

Table().with_columns('Simulated Statistic', simulated_statistics).hist()
plt.scatter(observed_statistic, -0.002, color='red', s=40);
```

---

**Question 1.5.** Using `simulated_statistics`, `observed_statistic`, and `num_simulations`, find the empirical p-value based on the simulation. **(4 points)**

_Hint:_ Reading [11.3.6](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#the-meaning-of-consistent) might be helpful for this question.


```python
p_value = ...
p_value
```

```python
grader.check("q1_5")
```

---

**Question 1.6.** Assign `correct_doctor` to the number corresponding to the correct statement below. Use the 5% cutoff for the p-value. **(4 points)**

1. The data support Dr. DeNero's position more than they support Dr. Sahai's.
2. The data support Dr. Sahai's position more than they support Dr. DeNero's.

As a reminder, here are the two claims made by Dr. DeNero and Dr. Sahai:
> **Doctor DeNero:** "Oh wow, it's just like tossing a coin that lands heads with chance $\frac{3}{5}$. Heads you get Vaccine 1 and Tails you get Vaccine 2."

>**Doctor Sahai:** "No, it's not. We're not doing anything like tossing a coin."


```python
correct_doctor = ...
correct_doctor
```

```python
grader.check("q1_6")
```

---

## You're done!

**Important submission information:**
- **Run all the tests** and verify that they all pass
- **Save** from the **File** menu
- **Run the final cell to grader.check_all()**
- **Click the download button to download the .ipynb file (see the figure below)**
- Then, go to Ulwazi and submit the .ipynb file to Homework 4: Testing a Model.

**It is your responsibility to make sure your work is saved before running the last cell.**


<img src="download_ipynb.png" alt="download ipynb" width="800"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order clicking the download button to download the .ipynb file. **Please save your work before downloading!**

<img src="download_ipynb.png" alt="download ipynb" width="800"/>

Tests for q1_1 (points: 6)
```
>>> type(sample_size) == int
True
```
```
>>> type(percent_V1) == float or type(percent_V1) == int
True
```
```
>>> 100 <= sample_size <= 500
True
```
```
>>> 1 <= percent_V1 <= 100
True
```


Tests for q1_2 (points: 5)
```
>>> type(valid_test_stat) == int
True
```
```
>>> any((valid_test_stat == x for x in np.arange(1, 5)))
True
```


Tests for q1_3 (points: 6)
```
>>> type(observed_statistic) == float
True
```
```
>>> 0 <= observed_statistic <= 100
True
```


Tests for q1_4 (points: 8)
```
>>> callable(one_simulated_statistic)
True
```
```
>>> type(simulated_statistics) == np.ndarray
True
```
```
>>> len(simulated_statistics) == num_simulations
True
```
```
>>> np.all(simulated_statistics >= 0)
True
```


Tests for q1_5 (points: 6)
```
>>> import numpy as np
>>> type(p_value) == float or type(p_value) == np.float64
True
```
```
>>> 0 <= p_value <= 1
True
```


Tests for q1_6 (points: 5)
```
>>> type(correct_doctor) == int
True
```
```
>>> any((correct_doctor == x for x in (1, 2)))
True
```
