# Homework 1: Arrays and Tables

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](https://ulwazi.wits.ac.za/courses/89081/pages/detailed-python-reference-sheet-python-cheat-sheet-2?module_item_id=1200407): Cheat sheet of helpful array & table methods used!

**Recommended Readings:**
- [Arrays](https://inferentialthinking.com/chapters/05/1/Arrays.html)
- [What is Data Science?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, **please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

```python
# Run this cell to set up the notebook, but please don't change it.

try:
    import pyodide_http
    pyodide_http.patch_all()
except ImportError:
    pass

import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("hw1.ipynb")
```

```python
# Run this cell and ignore the output. It is checking to see if you ran the first cell! 
# If you get a NameError, please run the cell at the very TOP of this notebook!

grader
```

---
---

## 1. Indexing Arrays

These exercises give you practice on accessing individual elements of arrays. In Python, each element is accessed by its index. An index indicates the position of an element within an array. For example, the first element is the element at index 0. Indices must be **integers**.

***Note:* If you have previous coding experience, you may be familiar with bracket notation. DO NOT use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.**

Be sure to refer to the [Python Reference](https://ulwazi.wits.ac.za/courses/89081/pages/detailed-python-reference-sheet-python-cheat-sheet-2?module_item_id=1200407) on the website if you feel stuck!

---

**Question 1.1.** The cell below creates an array of a few numbers.  Set `fourth_element` to the fourth element of `few_numbers`. **(3 Points)**

_**Hint:** Remember that the first element is at index 0, the second element is at index 1, and so on!_


```python
few_numbers = make_array(-2, -4, -7, -11, -16, -22)

fourth_element = ...
fourth_element
```

```python
grader.check("q1_1")
```

---

**Question 1.2.** The next cell creates a table that displays some information about the elements of `few_numbers` and their order. Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `gap_1`, `gap_2`, `gap_3`, and `gap_4` to the correct elements in the table. **(4 Points)**

*Hint:* Replace the `...` with strings or numbers. As a reminder, indices should be **integers**.


```python
gap_1 = ...
gap_2 = ...
gap_3 = ...
gap_4 = ...

elements_of_few_numbers = Table().with_columns(
    "English name for position", make_array(gap_1, "second", "third", "fourth", gap_2, "sixth"),
    "Index",                     make_array(0, gap_3, 2, 3, gap_4, 5),
    "Element",                   few_numbers)
elements_of_few_numbers
```

```python
grader.check("q1_2")
```

---

**Question 1.3.** The cell below loads an array called `president_birth_years`.  Calling `tbl.column(...)` on a table returns an array of the column specified, in this case the `Birth Year` column of the `president_births` table. The last element in that array is the most recent among the birth years of all the deceased Presidents. Assign that year to `most_recent_birth_year`. **(4 Points)**

**Note:** Avoid Googling the answer. You should be able to answer this question only using table methods.


```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q1_3")
```

---
---

## 2. Basic Array Arithmetic

---

**Question 2.1.** Multiply the numbers 42, -4224, 424224242, and 250 by 157. Assign each variable below such that `first_product` is assigned to the result of $42 * 157$, `second_product` is assigned to the result of $-4224 * 157$, and so on. **(3 Points)**

*Note*: For this question, **don't** use arrays.


```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print("First Product:", first_product)
print("Second Product:", second_product)
print("Third Product:", third_product)
print("Fourth Product:", fourth_product)
```

```python
grader.check("q2_1")
```

---

**Question 2.2.** Now, do the same calculation, but using an array called `numbers` and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`. **(3 Points)**


```python
numbers = ...
products = ...
products
```

```python
grader.check("q2_2")
```

---

**Question 2.3.** Oops, we made a typo!  Instead of 157, we wanted to multiply each number by 1577.  Compute the correct products in the cell below using array arithmetic.  Notice that your job is really easy if you previously defined an array containing the 4 numbers. **(3 Points)**


```python
correct_products = ...
correct_products
```

```python
grader.check("q2_3")
```

---

**Question 2.4.** We've loaded an array of temperatures in the next cell. Each number is the highest temperature observed on a day at a weather station in **Upington**, Northern Cape, between 1 January 2022 and 31 December 2025. The temperatures are in **degrees Celsius**.

A research group in the United States has asked for these figures in **Fahrenheit**, since that is the unit they work in. Convert all the temperatures to Fahrenheit by first multiplying them by $\frac{9}{5}$, then adding 32, i.e. $F = C * \frac{9}{5} + 32$. After converting, make sure to **ROUND** the final result to 2 decimal places using the `np.round` function. **(4 Points)**

```python
max_temperatures = Table.read_table("temperatures.csv").column("Max Daily Temperature")

max_temperatures_fahrenheit = ...
fahrenheit_temps_rounded = ...
fahrenheit_temps_rounded
```

```python
grader.check("q2_4")
```

---

**Question 2.5.** The cell below loads all the *lowest* temperatures from each day, also in degrees Celsius. Compute the daily temperature range for each day. That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature. **Pay attention to the units and give your answer in Fahrenheit!** Make sure **NOT** to round your answer for this question. **(3 Points)**

*Hint:* Use `min_temperatures` and/or `max_temperatures`, and be careful about *where* in your calculation you apply the conversion. A **difference** between two temperatures does not convert the same way a single temperature does - write the arithmetic out by hand if you are stuck.

*Check your work:* a typical daily range in Upington is 12-20 °C. Do your numbers look like **ranges**, or like **temperatures**?

```python
min_temperatures = Table.read_table("temperatures.csv").column("Min Daily Temperature")

fahrenheit_temperature_ranges = ...
fahrenheit_temperature_ranges
```

```python
grader.check("q2_5")
```

---
---

## 3. Old Faithful

Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

---

**Question 3.1.** The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct. **(4 Points)**


```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("q3_1")
```

---

**Question 3.2.** Assign `biggest_decrease` to the biggest decrease in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was 74 - 62 = 12 minutes. **(4 Points)**

*Hint*: We want to return the absolute value of the biggest decrease.

*Note*: `np.diff()` calculates the difference between subsequent values in an array. For example, calling `np.diff()` on the array `make_array(1, 8, 3, 5)` evaluates to `array([8 - 1, 3 - 8, 5 - 3])`, or `array([7, -5, 2])`.


```python
# np.diff() calculates the difference between subsequent values in a NumPy array
# where the first value is subtracted from the second value

differences = np.diff(waiting_times)
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("q3_2")
```

---

**Question 3.3.** Let’s imagine your guess for the next waiting time was always just the length of the previous waiting time. If you always guessed the previous waiting time, how big would your error in guessing the waiting times be, on average? **(4 Points)**

For example, since the first four waiting times are 79, 54, 74, and 62, the average difference between your guess and the actual time for just the second, third, and fourth eruptions would be $\frac{|79-54|+ |54-74|+ |74-62|}{3} = 19$.


```python
differences = np.diff(waiting_times)
average_error = ...
average_error
```

```python
grader.check("q3_3")
```

---
---

## 4. Tables

The questions in this section are about a roadside farm stall in the Mpumalanga Lowveld, which sells produce by the box.


---

**Question 4.1.** Suppose the stall has 4 crates of avocados, 3 crates of mangoes and 3 crates of naartjies. Create a table that contains this information. It should have two columns: `produce` and `crates`. Assign the new table to the variable `stock`. **(4 Points)**

*Note:* Use lower-case and singular words for the name of each kind of produce, like `"avocado"`.

*Note:* Review [Chapter 6](https://inferentialthinking.com/chapters/06/Tables.html) if you get stuck!

```python
# Our solution uses 1 statement split over 3 lines.
stock = ...
        ...
        ...
stock
```

```python
grader.check("q4_1")
```

---

**Question 4.2.** The file `stall_stock.csv` contains information about the stock at the farm stall. Each row represents the contents of one box of produce. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table. **(4 Points)**

```python
inventory = ...
inventory
```

```python
grader.check("q4_2")
```

---

**Question 4.3.** The file `saturday_sales.csv` contains the number of items sold from each box
last Saturday. It has an extra column called `price per item (R)` that gives the
price in rands of a *single item of produce* from that box. The rows are in the
same order as the `inventory` table.

Run the cell below to load it.

```python
# Just run this cell
sales = Table.read_table('saturday_sales.csv')
sales
```

How many items of produce did the stall sell in total on that day? **(5 Points)**

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q4_3")
```

---

**Question 4.4.** What was the stall's total revenue in rands (the total price of everything sold) on that day? **(5 Points)**

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the litchi sales.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q4_4")
```

---

**Question 4.5.** Make a new table called `remaining_inventory`. It should have the same rows and columns as `inventory`, except that the number of items sold from each box should be subtracted from that box's **original** count, so that the `count` column is **updated to be** the number of items remaining after Saturday. **(5 Points)**

*Hint:* `inventory` and `sales` have their rows in the same order, so the counts line up. You need to *replace* the `count` column, not add a second one.

```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("q4_5")
```

---
---

## 5. Differences Between Faculties

The University of the Witwatersrand publishes an annual *Facts and Figures* report. The table below shows the number of graduates from three Wits faculties in 2020 and in 2022.

| Faculty                                | 2020  | 2022  |
|----------------------------------------|-------|-------|
| Commerce, Law and Management           |  3059 |  3229 |
| Engineering and the Built Environment  |  1642 |  1530 |
| Health Sciences                        |  1373 |  1482 |


---

**Question 5.1.** For each faculty, compute the **absolute difference** between the number of graduates in 2020 and in 2022, and assign it to the name given below. Then assign `biggest_change` to the largest of the three.

Let Python perform all the arithmetic (like subtracting 3059 from 3229) rather than simplifying the expressions yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers. **(5 Points)**

```python
clm_change = ...
ebe_change = ...
health_change = ...

biggest_change = ...
biggest_change
```

```python
grader.check("q5_1")
```

---

**Question 5.2.** Which of the three faculties had the **smallest** absolute difference? Assign `smallest_change_faculty` to 1, 2, or 3 where each number corresponds to the following faculty:

1. Commerce, Law and Management
2. Engineering and the Built Environment
3. Health Sciences

Choose the number that corresponds to the faculty with the smallest absolute difference. **(4 Points)**

*Hint:* You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each faculty.

```python
smallest_change_faculty = ...
smallest_change_faculty
```

```python
grader.check("q5_2")
```

---

**Question 5.3.** For each faculty, define the "relative change" to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2020}} * 100}$

Fill in the code below such that `clm_relative_change`, `ebe_relative_change` and `health_relative_change` are assigned to the relative changes for their respective faculties.

You already computed the absolute differences in Question 5.1 — use those names rather than writing the subtractions out again. **(5 Points)**

```python
clm_relative_change = (clm_change / 3059) * 100
ebe_relative_change = ...
health_relative_change = ...
clm_relative_change, ebe_relative_change, health_relative_change
```

```python
grader.check("q5_3")
```

---

**Question 5.4.** Assign `biggest_rel_change_faculty` to 1, 2, or 3 where each number corresponds to the following:

1. Commerce, Law and Management
2. Engineering and the Built Environment
3. Health Sciences

Choose the number that corresponds to the faculty with the biggest relative change. **(4 Points)**

```python
biggest_rel_change_faculty = ...
biggest_rel_change_faculty
```

```python
grader.check("q5_4")
```

---

You're done with Homework 1!  

**Important submission information:** Be sure to run the tests and verify that they all pass, then choose **Save Notebook** from the **File** menu, then **run the final cell** and click the link to download the zip file. Then, go to Ulwazi and submit the zip file to the corresponding assignment. The name of this assignment is "DDDM Homework 1 Autograder". **It is your responsibility to make sure your work is saved before running the last cell.**

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

```python

```

Tests for q1_1 (points: 0)
```
>>> isinstance(fourth_element, (int, float, np.integer, np.floating))
True
```
```
>>> fourth_element < 0
True
```


Tests for q1_2 (points: 0)
```
>>> isinstance(gap_1, str) and isinstance(gap_2, str)
True
```
```
>>> isinstance(gap_3, int) and isinstance(gap_4, int)
True
```
```
>>> elements_of_few_numbers.num_rows == 6
True
```


Tests for q1_3 (points: 0)
```
>>> isinstance(most_recent_birth_year, (int, float, np.integer, np.floating))
True
```
```
>>> 1900 < most_recent_birth_year < 1970
True
```


Tests for q2_1 (points: 0)
```
>>> (first_product > 0, second_product < 0, third_product > 0, fourth_product > 0)
(True, True, True, True)
```


Tests for q2_2 (points: 0)
```
>>> len(products) == 4
True
```


Tests for q2_3 (points: 0)
```
>>> len(correct_products) == 4
True
```


Tests for q2_4 (points: 0)
```
>>> type(fahrenheit_temps_rounded) is np.ndarray
True
```
```
>>> len(fahrenheit_temps_rounded) == len(max_temperatures)
True
```
```
>>> np.all(fahrenheit_temps_rounded > 40)
True
```


Tests for q2_5 (points: 0)
```
>>> len(fahrenheit_temperature_ranges) == len(max_temperatures)
True
```
```
>>> np.all(fahrenheit_temperature_ranges > 0)
True
```
```
>>> np.all(fahrenheit_temperature_ranges < 55)
True
```


Tests for q3_1 (points: 0)
```
>>> 40 <= shortest <= 50
True
```
```
>>> shortest <= average <= longest
True
```


Tests for q3_2 (points: 0)
```
>>> isinstance(biggest_decrease, (int, float, np.integer, np.floating))
True
```
```
>>> biggest_decrease > 0
True
```


Tests for q3_3 (points: 0)
```
>>> 15 <= average_error <= 25
True
```


Tests for q4_1 (points: 0)
```
>>> isinstance(stock, tables.Table)
True
```
```
>>> stock.num_rows == 3 and stock.num_columns == 2
True
```


Tests for q4_2 (points: 0)
```
>>> isinstance(inventory, tables.Table)
True
```
```
>>> inventory.num_rows == 10
True
```


Tests for q4_3 (points: 0)
```
>>> isinstance(total_fruits_sold, (int, float, np.integer, np.floating))
True
```
```
>>> 400 < total_fruits_sold < 700
True
```


Tests for q4_4 (points: 0)
```
>>> isinstance(total_revenue, (int, float, np.integer, np.floating))
True
```
```
>>> 2000 < total_revenue < 5000
True
```


Tests for q4_5 (points: 0)
```
>>> isinstance(remaining_inventory, tables.Table)
True
```
```
>>> remaining_inventory.num_columns == inventory.num_columns
True
```
```
>>> np.all(remaining_inventory.column('count') >= 0)
True
```


Tests for q5_1 (points: 0)
```
>>> isinstance(biggest_change, (int, float, np.integer, np.floating))
True
```
```
>>> clm_change > 0 and ebe_change > 0 and health_change > 0
True
```
```
>>> 100 < biggest_change < 300
True
```


Tests for q5_2 (points: 0)
```
>>> smallest_change_faculty in (1, 2, 3)
True
```


Tests for q5_3 (points: 0)
```
>>> 0 < ebe_relative_change < 100
True
```
```
>>> 0 < health_relative_change < 100
True
```


Tests for q5_4 (points: 0)
```
>>> biggest_rel_change_faculty in (1, 2, 3)
True
```
