# Homework 3: Chance and Simulation

Read Chapters 9 and 10 of the textbook before starting.


Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
* [Python Reference](https://ulwazi.wits.ac.za/courses/89081/pages/detailed-python-reference-sheet-python-cheat-sheet-2?module_item_id=1200407)

**Recommended Readings**: 
* [Randomness](https://www.inferentialthinking.com/chapters/09/Randomness.html)
* [Sampling and Empirical Distributions](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

**This homework has hidden tests on it.** That means even though the tests may
say 100% passed, it doesn't mean your final grade will be 100%. We will be
running more hidden tests for correctness once everyone turns in the homework.

**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, or click the download button to download the .ipynb file. **Please save your work before downloading!**

<img src="download_ipynb.png" alt="download ipynb" width="800"/>

Let's begin by setting up the tests and imports by running the cell below.


## Do not change the code in the cell below ##

```python
# Run this cell to set up the notebook, but please don't change it.
%pip install -q urllib3<2.0 otter-grader datascience ipywidgets
try:
    import pyodide_http
    pyodide_http.patch_all()
except ImportError:
    pass

import numpy as np
from datascience import *

import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)

```

```python
# Run this cell to set up the notebook, but please don't change it.

import otter
grader = otter.Notebook("hw3.ipynb")
```

---
---

## 1. Roulette

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="roulette_wheel.jpeg" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

### Betting on Red ###
If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- If the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

---

**Question 1.1.** Define a function `dollar_bet_on_red` that takes the name of a color and returns your gain in dollars if that color had won and you had placed a one-dollar bet on red. Remember that the gain can be negative. Make sure your function returns an integer. **(4 points)**

*Note:* You can assume that the only colors that will be passed as arguments are red, black, and green. Your function doesn't have to check that.


```python
def dollar_bet_on_red(...):
    ...
```

```python
grader.check("q1_1")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

---

**Question 1.2.** Add a column labeled `Winnings: Red` to the table `wheel`. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. Your code should use the function `dollar_bet_on_red`. **(4 points)**

> *Hint:* You should not need a `for` loop for this question, instead try using a table method!

```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("q1_2")
```

### Simulating 10 Bets on Red
Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

--- 

**Question 1.3.** Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. Your table should have the same three column labels as in `wheel`. Once you've created that table, set `sum_bets` to your net gain in all 10 bets, assuming that you bet one dollar on red each time. **(4 points)**

*Note:* The [Sampling Methods Guide](https://drive.google.com/file/d/1k8uNN0Bg7J0BTgK-0Z5glDNtd_5FTJQl/view?usp=sharing) may be helpful!

*Hint:* It may be helpful to print out `ten_bets` after you create it! 

```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("q1_3")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

--- 

**Question 1.4.** Let's see what would happen if you made more bets. Define a function `net_gain_red` that takes the number of bets and returns the net gain (a single integer representing the total amount won from betting on red that many times). **(4 points)**

*Hint:* You should use your `wheel` table within your function.


```python
def net_gain_red(...):
    ...
```

```python
grader.check("q1_4")
```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_red(10)
```

---

**Question 1.5.** Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. After the cell is run, `simulated_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red. **(4 points)**

*Hint:* Think about which computational tool might be helpful for simulating a process multiple times. Section [9.3](https://inferentialthinking.com/chapters/09/3/Simulation.html) might be helpful!

*Note:* This cell might take a few seconds to run.


```python
num_bets = ...
repetitions = ...

simulated_gains_red = ...
...

len(simulated_gains_red) # Do not change this line! Check that simulated_gains_red is length 10000.
```

```python
grader.check("q1_5")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', simulated_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

---

**Question 1.6.** Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question. **(4 points)**


```python
loss_more_than_50 = True
```

```python
grader.check("q1_6")
```

---
---

## 2. Chances

Before you do this exercise, make sure you understand the logic behind all the examples in [Section 9.5](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen.
- Seeing if it is easier to find the chance that the event does not happen.

### Finding Chances

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green. In each part below, write an expression that evaluates to the chance of the event described.

---

**Question 2.1.** The color green wins **at least once** on the first 10 spins. **(5 points)**


```python
at_least_one_green = ...
```

```python
grader.check("q2_1")
```

--- 

**Question 2.2.** Two of the three colors **never** win in the first 10 spins. **(5 points)**

*Hint:* What situation(s) lead to two of the three colors never winning in the first 10 spins?


```python
lone_winners = ...
```

```python
grader.check("q2_2")
```

### Comparing Chances
In each of Questions 2.3 to 2.5, two events A and B are described. Choose from one of the following three options and set each answer variable to a single integer:

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time, and by thinking about the [law of averages](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html#the-law-of-averages).

---

**Question 2.3.** A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo. **(5 points)**

- Event A: all four different animals are picked (assuming the child picks without replacement)
- Event B: all four different animals are picked (assuming the child picks with replacement)


```python
toys_option = ...
```

```python
grader.check("q2_3")
```

---

**Question 2.4.** In a lottery, two numbers are drawn at random with replacement from the integers 1 through 1000. **(5 points)**

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


```python
lottery_option = ...
```

```python
grader.check("q2_4")
```

---

**Question 2.5.** A fair coin is tossed repeatedly. **(5 points)**

- Event A: There are 60 or more heads in 100 tosses
- Event B: There are 600 or more heads in 1000 tosses

*Hint*: Think about the law of averages!


```python
coin_option = ...
```

```python
grader.check("q2_5")
```

---
---

## 3. Three Ways Python Draws Random Samples

You have learned three ways to draw random samples using Python:

- `tbl.sample` draws a random sample of rows from the table `tbl`. The output is a table consisting of the sampled rows. 

- `np.random.choice` draws a random sample from a population whose elements are in an array. The output is an array consisting of the sampled elements.

- `sample_proportions` draws from a categorical distribution whose proportions are in an array. The output is an array consisting of the sampled proportions in all the categories. 

```python
# Look through this code and run this cell for questions 1 and 2
top_movies = Table.read_table('top_movies_2017.csv').select(0, 1)
top_movies.show(3)
```

```python
# Look through this code and run this cell for questions 1 and 2
studios_with_counts = top_movies.group('Studio').sort('count', descending=True)
studios_with_counts.show(3)
```

```python
# Look through this code and run this cell for questions 1 and 2
studios_of_all_movies = top_movies.column('Studio')
distinct_studios = studios_with_counts.column('Studio')

print("studios_of_all_movies:", studios_of_all_movies[:10], "...")
print("\n distinct_studios:", distinct_studios)
```

```python
# Look through this code and run this cell for questions 1 and 2
studio_counts_only = studios_with_counts.column('count')
studio_proportions_only = studio_counts_only / sum(studio_counts_only)

print("studio_counts_only:", studio_counts_only)
print("\n studio_proportions_only:", studio_proportions_only)
```

In Questions 1 and 2, we will present a scenario. For each scenario, we will ask whether the desired result can be achieved by using a given function and the following tables/arrays: `top_movies`, `studios_with_counts`, `studios_of_all_movies`, `distinct_studios`, `studio_counts_only` and `studio_proportions_only`. You can assume we know which index of the array corresponds with the studio with that movie count/proportion.

_Note: **Do not** explain your answer; please answer yes or no and the name of the array/table._

---

**Question 3.1.** Suppose you simulate a sample of 10 movies drawn at random with replacement from the 200 movies. Using **only that sample**, you want to output `True` if Paramount appears more often than Warner Brothers among the sampled studios, and `False` otherwise.

Which of the following would let you do this? Assign `sampling_method_1` to 1, 2, 3, or 4. **(6 points)**

1. Only `top_movies.sample(10)`
2. Only `np.random.choice(studios_of_all_movies, 10)`
3. `top_movies.sample(10)`, `np.random.choice(studios_of_all_movies, 10)` **and** `sample_proportions(10, studio_proportions_only)` — all three
4. None of them

*Hint:* Look carefully at what each one returns. Which of them keep enough information to count how often a particular studio appears?

```python
sampling_method_1 = ...
sampling_method_1
```

```python
grader.check("q3_1")
```

---

**Question 3.2.** Now suppose you want to output `True` if the **first** sampled movie was released by the same studio as the **last** sampled movie.

Assign `sampling_method_2` to 1, 2, 3, or 4. **(6 points)**

1. All three functions work, as in the previous question
2. `sample` and `np.random.choice` work, but `sample_proportions` does not
3. Only `sample_proportions` works
4. None of them work

*Hint:* `sample_proportions` tells you *how many* of the sample fell into each category. Does it tell you anything about the order they were drawn in?

```python
sampling_method_2 = ...
sampling_method_2
```

```python
grader.check("q3_2")
```

---

## You're done!

**Important submission information:**
- **Run all the tests** and verify that they all pass
- **Save** from the **File** menu
- **Run the final cell to generate the grader.check_all()**
- **Click the download button to download the .ipynb file (see the figure below)**
- Then, go to Ulwazi and submit the .ipynb file to Homework 3: Chance and Simulation.

**It is your responsibility to make sure your work is saved before running the last cell.**

Remember that this homework has hidden tests - passing every visible test does
not guarantee full marks.
<img src="download_ipynb.png" alt="download ipynb" width="800"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order clicking the download button to download the .ipynb file. **Please save your work before downloading!**

<img src="download_ipynb.png" alt="download ipynb" width="800"/>

Tests for q1_1 (points: 5)
```
>>> dollar_bet_on_red('red') == 1
True
```
```
>>> dollar_bet_on_red('black') == -1
True
```
```
>>> dollar_bet_on_red('green') == -1
True
```


Tests for q1_2 (points: 5)
```
>>> wheel.labels[2] == 'Winnings: Red'
True
```
```
>>> wheel.num_rows == 38
True
```


Tests for q1_3 (points: 5)
```
>>> set(['Pocket', 'Color', 'Winnings: Red']) == set(ten_bets.labels)
True
```
```
>>> ten_bets.num_rows == 10
True
```


Tests for q1_4 (points: 5)
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```


Tests for q1_5 (points: 5)
```
>>> type(simulated_gains_red) == np.ndarray
True
```
```
>>> len(simulated_gains_red) == 10000
True
```


Tests for q1_6 (points: 5)
```
>>> type(loss_more_than_50) == bool
True
```


Tests for q2_1 (points: 5)
```
>>> 0 < at_least_one_green < 1
True
```


Tests for q2_2 (points: 5)
```
>>> 0 < lone_winners < 1
True
```


Tests for q2_3 (points: 5)
```
>>> type(toys_option) == int
True
```
```
>>> 1 <= toys_option <= 3
True
```


Tests for q2_4 (points: 5)
```
>>> type(lottery_option) == int
True
```
```
>>> 1 <= lottery_option <= 3
True
```


Tests for q2_5 (points: 5)
```
>>> type(coin_option) == int
True
```
```
>>> 1 <= coin_option <= 3
True
```


Tests for q3_1 (points: 6)
```
>>> sampling_method_1 in (1, 2, 3, 4)
True
```


Tests for q3_2 (points: 6)
```
>>> sampling_method_2 in (1, 2, 3, 4)
True
```
